// variables4.rs
//
// Execute `rustlings hint variables4` or use the `hint` watch subcommand for a
// hint.

// I AM NOT DONE

fn main() {
    let mut x = 3;  // Ajoutez `mut` pour rendre x mutable
    println!("Number {}", x);
    x = 5;  // Maintenant, cette ligne est valide car x est mutable
    println!("Number {}", x);
}

